---

# Midterm Project
Ben Danzis(beda5109) and Taylor Thomas(tath3540)

#### Python skills
 - Iteration
 - Functions
 - GPT Models
 - Data structures
 - Reading in data
 - Asset Statements
 - Dictionary manipulation
 - Group By
 - Adding column to dataframe
 

# Hello, Colorado Marijuana Enforcement Division 

https://sbg.colorado.gov/marijuana-enforcement

Below is a snippet of Data from the Colorado Marijuana Enforcement Divison. 

In the excel doc below, there are a total of 60 entries. Each entry is a name of a person or business within the Colorado Marijuana Enforcement Divisons database.

Not every user has the legal authority to grow weed. Your midterm task is to parce through this dataset in order to find out who is NOT breaking any laws regarding the growth, and possible sale, of their weed. 

**Above each function will contain the necessary guidelines to write the function**

### Start off by reading in a dataset about weed growers in Colorado

- You will need to upload the file `COweed.xlsx`
- If you wish to learn more about marijuana, scroll down to the bottom of the notebook for a surprise

```python
import pandas as pd 
from pandas import DataFrame

filename = "COweed.xlsx"

df = pd.read_excel(filename)
```

<!-- grade_id: cell-cf3a441b96e0311a, points: 0, grade: false, solution: true -->
```python
def num_growers(input_df: DataFrame) -> int:
    '''
    Each row in a dataframe is called an "grower". This function should return the number of growers in the dataframe.
    
    Input:
    - input_df: a pandas DataFrame
    
    Output:
    - The number of growers in the dataframe
    - The output should be an integer
    '''
    return len(input_df)
print(num_growers(df))
    
```

Output:
```
60
```

<!-- grade_id: cell-3d6b3101ea84feba, points: 1, grade: true, solution: false -->
```python
assert type(num_growers(df)) == int, "num_observations should return an int"
```

<!-- grade_id: cell-dc63af484e5913aa, points: 0, grade: false, solution: true -->
```python
def customer_dataframe(input_df: DataFrame) -> DataFrame:
    '''
    return the same dataset as the input, but with the following columns:
    - Grower: the name of the grower
    - GrowerID: the unique ID of the grower
    - GrowerAge: the age of the grower
    - GrowerCity: the city of the grower
    - GrowerState: the state of the grower
    - GrowerStatus: the status of the grower
    - Jurisdictor: the jurisdiction of the grower
    - GrowerClassifcation: the classification of the grower
    - GrowerType: the type of the grower
    - GrowerDispensary: the dispensary of the grower
    - NumberofPlants: the number of plants of the grower
    - CostperPlant: the cost per plant of the grower
    '''
    # YOUR CODE HERE
    for col in input_df.columns:
        input_df.rename(columns={col:col.replace(" ","")}, inplace=True)
    return input_df
print(customer_dataframe(df))
    
```

Output:
```
growername  growerid  growerage  \
0                       Green Leaf Farms     10001         32   
1                     Rocky Mountain Bud     10002          4   
2                       Mohammed Hussein     10003         34   
3               High Altitude Herbalists     10004         99   
4                          Tegridy Farms     10005         69   
5                     Pikes Peak Pottery     10006         23   
6                       Carlos Fernandez     10007         21   
7                           Leila Rahman     10008         16   
8                Rocky Mountain Greenery     10009         44   
9                         Butters Stotch     10010         12   
10                            Luca Rossi     10011         90   
11                  Green Mountain Ganja     10012         80   
12                   Mile High Marijuana     10013         37   
13                           Fatima Khan     10014         33   
14                 Colorado Chronic Cure     10015     
```

```python

def customer_dictionary(input_df: DataFrame) -> dict:
    '''
    return a dictionary where the key is the GrowerID and the value is a dictionary of the other columns.
    '''
    # YOUR CODE HERE
    counter = 0
    customers = {}
    for index, row in input_df.iterrows():
        customers[counter] = row.to_dict()
        counter += 1
    return customers
print(customer_dictionary(df))
```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 1: {'growername': 'Rocky Mountain Bud', 'growerid': 10002, 'growerage': 4, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 500, 'costperplant': 200}, 2: {'growername': 'Mohammed Hussein', 'growerid': 10003, 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'c
```

```python
'''
Print the first row of the customers dictionary
'''
print(customer_dictionary(df)[59])
```

Output:
```
{'growername': 'Boulder Botanicals', 'growerid': 10060, 'growerage': 14, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'retail farm', 'numberofplants': 250, 'costperplant': 200}
```

### Inspecting the data

- In order to legally grow/sell weed there are multiple qualifications one has to pass
- In the functions below, we will be checking to see if the user has passed these qualifications
- If the user fails any of these qualifications, they will be removed from the dictionary

```python
def is_licensed(customer_dictionary: dict) -> dict:
    '''
    - The first function we will write is `is _licensed()`
    - This function will check to see if the user has a valid license
    - A 1 in the `growerstatus` column means the user has a valid license
    - A 0 in the `growerstatus` column means the user does not have a valid license
    - If the user has no license, they will be added to a new dictionary called `unlicensed`
    '''
    # YOUR CODE HERE
    for key, value in list(customer_dictionary.items()):
        if value['growerstatus'] == 0:
            del customer_dictionary[key]
    return customer_dictionary
print(is_licensed(customer_dictionary(df)))
```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 1: {'growername': 'Rocky Mountain Bud', 'growerid': 10002, 'growerage': 4, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 500, 'costperplant': 200}, 2: {'growername': 'Mohammed Hussein', 'growerid': 10003, 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'c
```

```python
assert type(is_licensed(customer_dictionary(df))) == dict, "is_licensed should return a dictionary"
```

<!-- grade_id: cell-4ce832836aeee7b4, points: 0, grade: false, solution: true -->
```python
def of_age(customer_dictionary: dict) -> dict:
    '''
    - The next function we will write is `of_age()`
    - This function will check to see if the user is at least 21 years old
    - Check age by looking at the `growerage` column
    - If the user is under 21, they will be removed from the dictionary
    '''
    # YOUR CODE HERE
    customer_dictionary = is_licensed(customer_dictionary)
    for key, value in list(customer_dictionary.items()):
        if value['growerage'] < 21:
            del customer_dictionary[key]
    return customer_dictionary
print(of_age(customer_dictionary(df)))

```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 2: {'growername': 'Mohammed Hussein', 'growerid': 10003, 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'costperplant': 100}, 3: {'growername': 'High Altitude Herbalists', 'growerid': 10004, 'growerage': 99, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail store', 'numberofplants
```

```python
''' 
Print the number of customers in the `customer_dictionary` after running `is_licensed()` and `of_age()`
'''
print(len(of_age(is_licensed(customer_dictionary(df)))))
```

Output:
```
38
```

```python
assert len(of_age(is_licensed(customer_dictionary(df)))) == 38, "There should be 38 customers in the dictionary after running is_licensed() and of_age()"
```

<!-- grade_id: cell-7b66971d672c06e1, points: 0, grade: false, solution: true -->
```python
def group_by_growertype_and_city(input_df: DataFrame) -> DataFrame:
    '''
    
    In this function you will practice grouping data
    and applying an aggregation. To complete the function
    
    - Take an input data frame
    - Group it by growertype
    - Call the "count" aggregation function to count up the number of entity types in 
      each group
    
    Input:
    - input_df: a pandas DataFrame
    
    Output:
    - output_df: a pandas DataFrame, with data grouped by entity type showing the
                 count of entities in each group
                 
    Call the reset index at the end
    '''
    # YOUR CODE HERE
    output_df = input_df.groupby(['growertype', 'growercity']).size().reset_index(name='counts')
    return output_df
print(group_by_growertype_and_city(df))
    
```

Output:
```
growertype         growercity  counts
0       medical               Aspen       2
1       medical              Aurora       1
2       medical             Boulder       4
3       medical        Breckenridge       1
4       medical    Colorado Springs       4
5       medical              Denver       4
6       medical              Pueblo       2
7       medical                Vail       3
8       personal            Alamosa       2
9       personal             Aurora       1
10      personal            Boulder       1
11      personal         Broomfield       1
12      personal        Castle Rock       1
13      personal               Erie       1
14      personal           Fredrick       1
15      personal    Grand Junction        1
16      personal             Greely       1
17      personal           Lakewood       1
18      personal              Rifle       1
19      personal         South Park       2
20      personal  Steamboat Springs       1
21      personal          Telluride   
```

### Now things get challenging (challenge yourself! not optional >:)
- If you have made it this far, you are doing great!
- You have converted a Dataframe to a customer Dictionary
- We have ruled out all growers either under age or unlicensed and updated the dictionary
- Now we have to seperate each grower into their own growertype
- There are 3 growertypes: Medical, Retail, and Personal
- Each growertype has different limits to the amount of weed they can grow

To do the challenge, you will need to split the updated customer dictionary into 3 different dictionaries. Each dictionary will contain the information for each growertype. Further farming is required before we can create the final customer DataFrame.

https://www.365cannabis.com/colorado-cultivation-license-and-marijuana-laws-2/#:~:text=A%20Colorado%20Medical%20Marijuana%20Cultivation,3%20(1501%2D3000%20plants)

```python
def split_by_growertype(customer_dictionary: dict) -> dict:
    '''
    - First write a function that splits the updated dictionary into three new dictionaries:
        - `meddict`: contains all medical growers
        - `recdict`: contains all recreational growers
        - `persdict`: contains all personal growers
    - Each grower type can be found in the `growertype` column
    '''
    customer_dictionary = of_age(customer_dictionary)
    meddict = {}
    recdict = {}
    persdict = {}
    for key, value in list(customer_dictionary.items()):
        if value['growertype'] == 'medical ':
            meddict[key] = value
        elif value['growertype'] == 'recreational':
            recdict[key] = value
        elif value['growertype'] == 'personal':
            persdict[key] = value
    return meddict, recdict, persdict

```

```python
'''Print the new medical dictionary and number of medical growers'''
print(split_by_growertype(customer_dictionary(df))[0])
print(len(split_by_growertype(customer_dictionary(df))[0]))

```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 11: {'growername': 'Green Mountain Ganja', 'growerid': 10012, 'growerage': 80, 'growercity': 'Vail', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 450, 'costperplant': 200}, 12: {'growername': 'Mile High Marijuana', 'growerid': 10013, 'growerage': 37, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 3
```

```python
'''Print the new recreational dictionary and number of recreational growers'''
print(split_by_growertype(of_age(is_licensed(customer_dictionary(df))))[1])
print(len(split_by_growertype(of_age(is_licensed(customer_dictionary(df))))[1]))
```

Output:
```
{3: {'growername': 'High Altitude Herbalists', 'growerid': 10004, 'growerage': 99, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail store', 'numberofplants': 1000, 'costperplant': 100}, 4: {'growername': 'Tegridy Farms', 'growerid': 10005, 'growerage': 69, 'growercity': 'South Park', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail farm', 'numberofplants': 1600, 'costperplant': 100}, 5: {'growername': 'Pikes Peak Pottery', 'growerid': 10006, 'growerage': 23, 'growercity': 'Colorado Springs', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail 
```

```python
'''Print the new personal dictionary and number of personal growers'''
print(split_by_growertype(of_age(is_licensed(customer_dictionary(df))))[2])
print(len(split_by_growertype(of_age(is_licensed(customer_dictionary(df))))[2]))
```

Output:
```
{2: {'growername': 'Mohammed Hussein', 'growerid': 10003, 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'costperplant': 100}, 10: {'growername': 'Luca Rossi', 'growerid': 10011, 'growerage': 90, 'growercity': 'Lakewood', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'costperplant': 100}, 19: {'growername': 'Lina Chen', 'growerid': 10020, 'growerage': 40, 'growercity': 'Castle Rock', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 9, 'costperplant': 1
```

```python
def medical_growers(meddict: dict) -> dict:
    '''
    - The first growertype is "Medical"
    - Each medical grower can grow up to 500 plants per year (Colorado Marijuana Enforcement Division Cap)
    - If they grow more than the limit, remove them from the dictionary
    '''
    for key, value in list(meddict.items()):
        if value['numberofplants'] > 500:
            del meddict[key]
    return meddict
print(medical_growers(split_by_growertype(customer_dictionary(df))[0]))

```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 11: {'growername': 'Green Mountain Ganja', 'growerid': 10012, 'growerage': 80, 'growercity': 'Vail', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 450, 'costperplant': 200}, 12: {'growername': 'Mile High Marijuana', 'growerid': 10013, 'growerage': 37, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 3
```

```python
assert type(medical_growers(split_by_growertype(customer_dictionary(df))[0])) == dict, "medical_growers should return a dictionary"
```

```python
def recreational_growers(recdict: dict) -> dict:
    '''
    - The first growertype is "Recreational"
    - Each medical grower can grow up to 1800 plants per year (Colorado Marijuana Enforcement Division Cap)
    - If they grow more than the limit, remove them from the dictionary
    '''
    for key, value in list(recdict.items()):
        if value['numberofplants'] > 1800:
            del recdict[key]
    return recdict
print(recreational_growers(split_by_growertype(customer_dictionary(df))[1]))
```

Output:
```
{3: {'growername': 'High Altitude Herbalists', 'growerid': 10004, 'growerage': 99, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail store', 'numberofplants': 1000, 'costperplant': 100}, 4: {'growername': 'Tegridy Farms', 'growerid': 10005, 'growerage': 69, 'growercity': 'South Park', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail farm', 'numberofplants': 1600, 'costperplant': 100}, 17: {'growername': 'Summit Sensimilla', 'growerid': 10018, 'growerage': 34, 'growercity': 'Breckenridge', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'recreational', 'growerydispensary': 'retail stor
```

```python
def personal_growers(persdict: dict) -> dict:
    '''
    - The first growertype is "Personal"
    - Each medical grower can grow up to 6 plants per year (Colorado Marijuana Enforcement Division Cap)
    - Technically, adults living in the same household can grow up to 6 plants each, up to 36 plants, 
        but each adult has to be registered personally with the state, so 6 is the limit
    - If they grow more than the limit, remove them from the dictionary
    '''
    for key, value in list(persdict.items()):
        if value['numberofplants'] > 6:
            del persdict[key]
    return persdict
print(personal_growers(split_by_growertype(customer_dictionary(df))[2]))
```

Output:
```
{2: {'growername': 'Mohammed Hussein', 'growerid': 10003, 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'costperplant': 100}, 10: {'growername': 'Luca Rossi', 'growerid': 10011, 'growerage': 90, 'growercity': 'Lakewood', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 6, 'costperplant': 100}, 35: {'growername': 'Noura Al-Farsi', 'growerid': 10036, 'growerage': 48, 'growercity': 'Rifle', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': 'personal', 'growerydispensary': 'personal', 'numberofplants': 3, 'costperplant': 10
```

```python
'''Count the number of growers still in each dictionary'''
meddict = medical_growers(split_by_growertype(customer_dictionary(df))[0])
recdict = recreational_growers(split_by_growertype(customer_dictionary(df))[1])
persdict = personal_growers(split_by_growertype(customer_dictionary(df))[2])
len(meddict) + len(recdict) + len(persdict)
```

Output:
```
34
```

```python
assert len(meddict) + len(recdict) + len(persdict) == 34, "There should be 34 growers in the dictionaries after running all the functions"
```

#### Final Stretch (not optional)
- Still alive? Great! Now we have to remerege the 3 dictionaries into one dictionary
- At this point, the remainging growers are all legal (hopefully you have 34 growers left)
- Once you have remereged the dictionaries, you will need to do some math
- Money talks, find out the total amount of money each grower owes
- Create a new column in the customer DataFrame called `totalmoneyowed`
- Final step, print the finalized customer DataFrame for the Colorado Marijuana Enforcement Division

```python
'''
- The first step is to remerge the three dictionaries into one dictionary
- Call the final dictionary `finaldict`
'''
finaldict = {**meddict, **recdict, **persdict}
print(finaldict)
```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 11: {'growername': 'Green Mountain Ganja', 'growerid': 10012, 'growerage': 80, 'growercity': 'Vail', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 450, 'costperplant': 200}, 12: {'growername': 'Mile High Marijuana', 'growerid': 10013, 'growerage': 37, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 3
```

```python
len(finaldict)
```

Output:
```
34
```

```python
assert len(finaldict) == len(meddict) + len(recdict) + len(persdict), "The final dictionary should have the same number of growers as the three dictionaries combined"
```

https://apotforpot.com/blogs/growing/cannabis-costs/

```python
def total_money_owed(finaldict: dict) -> dict:
    '''
    - Write a function that calculates the total money owed for the year
    - The function should take the `finaldict` as an input
    '''
    for key, value in finaldict.items():
        value['totalmoneyowed'] = value['numberofplants'] * value['costperplant']
    return finaldict
print(total_money_owed(finaldict))
```

Output:
```
{0: {'growername': 'Green Leaf Farms', 'growerid': 10001, 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200, 'totalmoneyowed': 100000}, 11: {'growername': 'Green Mountain Ganja', 'growerid': 10012, 'growerage': 80, 'growercity': 'Vail', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 450, 'costperplant': 200, 'totalmoneyowed': 90000}, 12: {'growername': 'Mile High Marijuana', 'growerid': 10013, 'growerage': 37, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 1, 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'grow
```

```python
'''
- Are you ready to be done with this project? One final step...
- Turn your ditctionary into a pandas DataFrame
'''
finaldf = pd.DataFrame.from_dict(finaldict, orient='index')
print(finaldf)
```

Output:
```
growername  growerid  growerage  \
0                       Green Leaf Farms     10001         32   
11                  Green Mountain Ganja     10012         80   
12                   Mile High Marijuana     10013         37   
14                 Colorado Chronic Cure     10015         42   
27                    Breck Bud Brothers     10028         36   
32                 Boulder Buds Boutique     10033         57   
36  Colorado Springs Cannabis Collective     10037         59   
41                  Pueblo Pot Producers     10042         43   
43                   High Altitude Herbs     10044         28   
44                  Mile High Medicinals     10045         41   
50               Rocky Mountain Remedies     10051         29   
53                     Pueblo Pot Palace     10054         85   
54                Denver Dope Dispensary     10055         56   
56                     Vail Valley Vapor     10057         93   
3               High Altitude Herbalists     10004     
```

### Bonus + insights & conclusions
- If you have made it this far, you are a true weed enforcer
- You have created a dictionary of legal growers
- You have seperated the growers into their own growertype
- You have remereged the growers into one dictionary
- You have calculated the total amount of money each grower owes
- You have created a new column in the customer DataFrame called `totalmoneyowed`
- You have printed the finalized customer DataFrame for the Colorado Marijuana Enforcement Division
- You have completed the bonus
- Now take the time to ask the encoded GPT model a question about weed

## Conclusions and Insights
Throughout the course of this project we have gained a wealth of knowledge about utilizing
Python to describe business processes, as well as acquiring intriguing specifications about the
marijuana industry in Colorado. Our primary focus was to examine the various qualifications and
characteristics of different growers in the region We took into consideration if the grower was
licensed or not, if the grower is of age (over twenty-one years), the city the grower resides in,
and the type of grower they are (recreational, personal, or medicinal). Python allowed us to
break down the business process into smaller components.
    The first component we analyze is whether or not the grower is a licensed grower or not. In
order to legally grow marijuana in Colorado, you must have a license. The price depends on
what type of grower you are (recreational, personal, or medicinal). Using Python, we were able
to accurately identify and remove those who were not licensed to ensure that only authorized
growers are included in the dictionary. We go on to perform the same analysis to test whether or
not the grower is over twenty-one years old. Later on, we split this dictionary into three separate
dictionaries for recreational, personal, and medicinal growers. By doing this, we can gain
insights into the different needs and preferences of each group which can essentially aid in
developing sales strategies and product development. The output dictionaries can also be used
to perform further analysis on each customer group by, for example, identifying the most popular
products among each group.
    We proceed to group the data by grower type and grower city. By doing so, we can gain insights
into what cities growers are located and how they are distributed across grower types. This
information could later be useful for targeted marketing campaigns, tracking grower behavior, or
for identifying areas where the grower may need to expand its services.
Next, we filtered all three separate dictionaries (personal, recreational, and medical) by the
number of plants they are allowed to grow to make sure each grower is in compliance with this
number, which was determined by the Colorado Marijuana Enforcement Division. Any grower
who violated the limit was removed from their respective dictionary. This code ensures
compliance with regulations and helps growers to avoid potential legal issues so that they can
continue growing marijuana happily and legally.
    The final dictionary provides a list of all growers who are in compliance with the regulations and
specifications that make their growing legal. Additionally, we include a function that calculates
the total amount of money owed for each grower. This is useful for keeping track of each
growers expenses and revenue. Growers can use this data to analyze their sales and see
where they stand in terms of financial health to aid in business and process planning.
    Python was a valuable tool as it allowed us to easily manipulate and analyze large sets of data,
breaking down the business processes into smaller components and providing valuable insights
into the marijuana industry in Colorado.

```python
import openai
openai.api_key = "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"
```

```python
messages = [{"role": "system", 
              "content": "What is the classic phrase used at the start of many programming assignments?."}]

completion = openai.ChatCompletion.create(
  model="gpt-3.5-turbo",
  messages=messages
)
```

```python
import time

def finding_answers(prompt: str) -> str:
    
    time.sleep(5)
    
    messages = [{"role": "system",
                "content": prompt}]
    completion = openai.ChatCompletion.create(
        model="gpt-3.5-turbo",
        messages=messages
    )
    return str(completion["choices"][0]["message"]["content"])
print(finding_answers("What are health benefits of marijuana?"))
```

Output:
```
As an AI language model, I don't have personal opinions or beliefs, but the following are some of the possible health benefits of marijuana, as per some scientific studies: 

1. Pain relief: Marijuana is widely known for its pain-relieving properties. It contains cannabinoids that have shown to be effective in reducing pain and inflammation, especially in chronic conditions like arthritis and multiple sclerosis.

2. Reducing anxiety and depression: Marijuana has shown to reduce symptoms of anxiety and depression, improve mood, and enhance relaxation in some individuals.

3. Reducing seizures: There is growing evidence that CBD, a cannabinoid found in marijuana, can help reduce seizures in people with epilepsy.

4. Improve sleep: Marijuana has been found to improve the quality of sleep and reduce the time it takes to fall asleep.

5. Reducing symptoms of cancer treatment: Marijuana has been found to reduce nausea and vomiting caused by chemotherapy and improve appetite in people undergo
```