---

# Midterm Project
Ben Danzis and Taylor Thomas

#### Python skills
 - V

# Hello, Colorado Marijuana Enforcement Division 

Below is a snippet of Data from the Colorado Marijuana Enforcement Divison. 

In the excel doc below, there are a total of 60 entries. Each entry is a name of a person or business within the Colorado Marijuana Enforcement Divisons database.

Not every user has the legal authority to grow weed. Part of our midterm task was to parce through this dataset in order to find out who is NOT breaking any laws regarding the growth, and possible sale, of their weed. 

**Above each function will contain the necessary guidelines to write the function**

### Start off by reading in a dataset about weed growers in Colorado

- You will need to upload the file `COweed.csv`

```python
import pandas as pd 
from pandas import DataFrame

filename = "COweed.xlsx"

df = pd.read_excel(filename)
```

<!-- grade_id: cell-cf3a441b96e0311a, points: 0, grade: false, solution: true -->
```python
def num_growers(input_df: DataFrame) -> int:
    '''
    Each row in a dataframe is called an "grower". This function should return the number of growers in the dataframe.
    
    Input:
    - input_df: a pandas DataFrame
    
    Output:
    - The number of growers in the dataframe
    - The output should be an integer
    '''
    return len(input_df)
print(num_growers(df))
    
```

Output:
```
60
```

<!-- grade_id: cell-3d6b3101ea84feba, points: 1, grade: true, solution: false -->
```python
assert type(num_growers(df)) == int, "num_observations should return an int"
```

<!-- grade_id: cell-dc63af484e5913aa, points: 0, grade: false, solution: true -->
```python
def customer_dictionary(input_df: DataFrame) -> DataFrame:
    '''
    return the same dataset as the input, but with the following columns:
    - Grower: the name of the grower
    - GrowerID: the unique ID of the grower
    - GrowerAge: the age of the grower
    - GrowerCity: the city of the grower
    - GrowerState: the state of the grower
    - GrowerStatus: the status of the grower
    - Jurisdictor: the jurisdiction of the grower
    - GrowerClassifcation: the classification of the grower
    - GrowerType: the type of the grower
    - GrowerDispensary: the dispensary of the grower
    - NumberofPlants: the number of plants of the grower
    - CostperPlant: the cost per plant of the grower
    '''
    # YOUR CODE HERE
    for col in input_df.columns:
        input_df.rename(columns={col:col.replace(" ","")}, inplace=True)
    return input_df
print(customer_dictionary(df))
    
```

Output:
```
growerid                            growername  growerage  \
0      10001                      Green Leaf Farms         32   
1      10002                    Rocky Mountain Bud         45   
2      10003                      Mohammed Hussein         34   
3      10004              High Altitude Herbalists         99   
4      10005                         Tegridy Farms         69   
5      10006                    Pikes Peak Pottery         23   
6      10007                      Carlos Fernandez         21   
7      10008                          Leila Rahman         16   
8      10009               Rocky Mountain Greenery         44   
9      10010                        Butters Stotch         12   
10     10011                            Luca Rossi         90   
11     10012                  Green Mountain Ganja         80   
12     10013                   Mile High Marijuana         37   
13     10014                           Fatima Khan         33   
14     10015                 
```

```python
def dataframe_to_dictionary(input_df: DataFrame) -> dict:
    '''
    Convert the dataframe to a dictionary, where each grower is a key, 
    and the value is a dictionary of the grower's information.
    - Each key should be on its own line when printed
    '''
    # YOUR CODE HERE
    return input_df.to_dict(orient='index')
print(dataframe_to_dictionary(df))
```

Output:
```
{0: {'growerid': 10001, 'growername': 'Green Leaf Farms', 'growerage': 32, 'growercity': 'Denver', 'growerstate': 'CO', 'growerstatus': 'licensed', 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical farm', 'numberofplants': 500, 'costperplant': 200}, 1: {'growerid': 10002, 'growername': 'Rocky Mountain Bud', 'growerage': 45, 'growercity': 'Boulder', 'growerstate': 'CO', 'growerstatus': 'licensed', 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'profit', 'growertype': 'medical ', 'growerydispensary': 'medical store', 'numberofplants': 500, 'costperplant': 200}, 2: {'growerid': 10003, 'growername': 'Mohammed Hussein', 'growerage': 34, 'growercity': 'Telluride', 'growerstate': 'CO', 'growerstatus': 'licensed', 'jurisdictor': 'Colorado Marijuana Enforcement Division ', 'growerclassification': 'personal', 'growertype': nan, 'growerydispensary': nan, 'numberof
```

### Inspecting the data

<!-- grade_id: cell-4ce832836aeee7b4, points: 0, grade: false, solution: true -->
```python
def num_entity_types(input_df: DataFrame) -> int:
    '''
    
    The pd.unique method returns the number of unique items
    in a pandas series.
    
    Use the unique method to get the total number of distinct 
    entity types in the dataframe
    
    Input:
    - input_df: a pandas DataFrame
    
    Output:
    - The number of total unique entity types in the 
      dataset.
    '''
    # YOUR CODE HERE
    return len(pd.unique(input_df['entitytype']))
print(num_entity_types(df))
```

Output:
```
25
```

<!-- grade_id: cell-59e427abe727e1b4, points: 2, grade: true, solution: false -->
```python
# this is a two point hidden test.
# it checks that your `num_entity_types` is correct
# for a subset of the data

```

<!-- grade_id: cell-7b66971d672c06e1, points: 0, grade: false, solution: true -->
```python
def group_by_entity_type_and_count(input_df: DataFrame) -> DataFrame:
    '''
    
    In this function you will practice grouping data
    and applying an aggregation. To complete the function
    
    - Take an input data frame
    - Group it by entity type
    - Call the "count" aggregation function to count up the number of entity types in 
      each group
    
    Input:
    - input_df: a pandas DataFrame
    
    Output:
    - output_df: a pandas DataFrame, with data grouped by entity type showing the
                 count of entities in each group
                 
    You can call reset_index if you want, but you don't have to
    '''
    # YOUR CODE HERE
    return input_df.groupby("entitytype").count()
print(group_by_entity_type_and_count(df))
```

Output:
```
entityid  entityname  \
entitytype                                                                 
Cooperative                                               28          28   
Cooperative Association                                    7           7   
Cooperative Association-Public Benefit Corporation         3           3   
Cooperative-Public Benefit Corporation                     8           8   
Corporation                                            15878       15878   
Corporation Sole                                           8           8   
Corporation-Public Benefit Corporation                    98          98   
Foreign Cooperative                                        3           3   
Foreign Corporation                                     1779        1779   
Foreign Limited Liability Company                       1948        1948   
Foreign Limited Liability Limited Partnership             17          17   
Foreign Limited Liability Partnership                     10    
```

<!-- grade_id: cell-7152594f79fe4839, points: 1, grade: true, solution: false -->
```python
# this is a one point test

# call the function
type_and_count = group_by_entity_type_and_count(df)

# reset the index
reset = type_and_count.reset_index().set_index("entitytype")

# check that there are 15878 corporations in the result
assert reset.loc["Corporation"]["entityid"] == 15878, "There are 15878 total corporations registered in Boulder"
```

<!-- grade_id: cell-aa0cb4813000c99d, points: 2, grade: true, solution: false -->
```python
# this is a two point hidden test

# The test checks if your group_by_entity_type_and_count function
# can also work for an unseen dataframe.
# To pass this test, check your own code to ensure that your function
# works for different inputs

```

### Challenge problem
- This problem is worth one point
- It is a little harder than some of the other ones in this notebook. You will have to do a little bit of research on the pandas API on your own to solve the problem.
- If you get everything right but the challenge you will get a 9/10 on the HW

To do the challenge, you need to make sure the entityformdate is correctly represented as a date to the problem. Just run the line of code below to ensure the date is correctly formatted.

```python
df["entityformdate"] = pd.to_datetime(df["entityformdate"])
```

Ok! Now you are ready for the challenge problem. Using groupby and consulting the [pandas documentation](https://pandas.pydata.org/docs/user_guide/groupby.html#named-aggregation), write a function that takes a dataframe and entity type as input and returns the earliest formation date for the businesses in the group. For instance, the oldest "Nonprofit Corporation" in Boulder (according to this dataset) was formed on 1871-03-03. So you should return 1871-03-03 (a Timestamp) if the input was this dataframe and the entity type "Nonprofit Corporation". Timestamps are just pandas way of representing dates; see the docs below. 

#### Learn to read the docs!! 

Working from documentation is a *really* important skill. It's not like any class can teach you every feature of a language or library. A class like BAIM 3220 can get you up and running. But you will need to consult documentation to learn how to use new features as you use Python (or pandas) in the future.

<!-- grade_id: cell-b15a6503ded7b004, points: 0, grade: false, solution: true -->
```python
from pandas import Timestamp
# see https://pandas.pydata.org/docs/reference/api/pandas.Timestamp.html

def challenge(input_dataframe: DataFrame, entitytype: str) -> Timestamp:
    # YOUR CODE HERE
    return df[df["entitytype"] == entitytype]["entityformdate"].min()
print(challenge(df, "Corporation"))
```

Output:
```
1894-09-08 00:00:00
```

```python
# this is an ungraded test to show you the expected outcome from the function. The hidden test is very similar
assert df.groupby("entitytype")["entityformdate"].min()["Nonprofit Corporation"].year == 1871
```

<!-- grade_id: cell-f8a7d64b004be39f, points: 1, grade: true, solution: false -->
```python
# this is a 1 point hidden test

```